# Descriptive Stats

For the questions below, we're going to refer exclusively to an HHCAPS survey data set that is available in `/data/hhcaps.csv`.  Use whatever commands you want to calculate the information required to get to the answer.


Put your solution as the last line right before the test cell, using the form shown below for computing the number of rows in a dataframe.

```
answer = my_df.shape[0]
```

The assertions will often give you what the final answer should be, but you won't receive any points unless you compute the answer using code.  For instance, if you just typed `answer = 132` for the numer of rows instead of `answer = my_df.shape[0]`, you would not receive credit.


```python
%matplotlib inline
import pandas as pd

hhcaps = pd.read_csv('/data/hhcaps.csv')
hhcaps
```

Output:
```
State  CMS Certification Number (CCN)*  \
0        AL                            17000   
1        AL                            17008   
2        AL                            17009   
3        AL                            17013   
4        AL                            17014   
...     ...                              ...   
11995    TX                           747978   
11996    TX                           747979   
11997    TX                           747980   
11998    TX                           747981   
11999    TX                           747982   

                                           Provider Name  \
0      BUREAU OF HOME & COMMUNITY SERVICES           ...   
1      JEFFERSON COUNTY HOME CARE                    ...   
2      ALACARE HOME HEALTH & HOSPICE                 ...   
3      GENTIVA HEALTH SERVICES                       ...   
4      AMEDISYS HOME HEALTH                          ...   
...                                                  ...   
11995  EX
```

## #01 - 

How many columns does this file contain?

```python
hhcaps.shape
```

Output:
```
(12000, 39)
```

<!-- grade_id: q01-solution, points: 0, grade: false, solution: true -->
```python
answer = hhcaps.shape[1]
```

<!-- grade_id: q01-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 39)
```

```python

```

## #02 - 

How many different values for State are there in this file?

<!-- grade_id: q02-solution, points: 0, grade: false, solution: true -->
```python
states = hhcaps.groupby('State')
answer = len(states)


```

<!-- grade_id: q02-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 55)
```

```python
print(answer)
```

Output:
```
55
```

## #03 - 

Which of those State values has the highest frequency of occurence?

<!-- grade_id: q03-solution, points: 0, grade: false, solution: true -->
```python
state_numbers = hhcaps.groupby('State').count()['Provider Name'].sort_values()
answer = state_numbers.index[len(state_numbers) -1]
```

<!-- grade_id: q03-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'TX')
```

```python
print(answer)
```

Output:
```
TX
```

## #04 - 

Which of those State values has the best average performance on the `Star Rating for health team communicated well with them` score?

<!-- grade_id: q04-solution, points: 0, grade: false, solution: true -->
```python
starcom_score = states.agg('mean')['Star Rating for health team communicated well with them'].sort_values(ascending=False)
answer = starcom_score.index[0]

print(answer)
```

Output:
```
ME
```

<!-- grade_id: q04-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'ME')
```

```python

```

## #05 - 

What was the average score on `Star Rating for how patients rated overall care from agency` for providers listed as having a `Type of Ownership` of `Hospital Based Program`

<!-- grade_id: q05-solution, points: 0, grade: false, solution: true -->
```python
hospital_program = hhcaps[hhcaps['Type of Ownership'] == "Hospital Based Program"]
answer = hospital_program.agg('mean')['Star Rating for how patients rated overall care from agency']

print(answer)
```

Output:
```
3.6925207756232687
```

<!-- grade_id: q05-test, points: 2, grade: true, solution: false -->
```python
import numpy
assert(type(answer) == float or type(answer) == numpy.float64)
assert(round(answer,5) == round(3.6925207756232687,5))

```

```python

```

## #06 - 

There is one column in the file that has the same value on every row.  What is the name of that column?

```python
answer = str(hhcaps.nunique().sort_values().index[0])
```

<!-- grade_id: q06-solution, points: 0, grade: false, solution: true -->
```python
print(answer)
```

Output:
```
Offers Nursing Care Services
```

<!-- grade_id: q06-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer.lower() == 'Offers Nursing Care Services'.lower())

```

```python

```

## #07 - 

There are six (6) columns in the file that indicate (True or False) if the provider offers certain services.  Those column names all start with `Offers...`.  How many facilities offer **everything** that the survey was interested in asking about?  That is, how many rows have True in all six of those columns?

<!-- grade_id: q07-solution, points: 0, grade: false, solution: true -->
```python
offers = hhcaps.select_dtypes(include='bool')
all_offers = offers[(offers.all(axis='columns') == True)]
answer = len(all_offers)
```

<!-- grade_id: q07-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 9029)

```

```python
print(answer)
```

Output:
```
9029
```

## #08 - 

The survey results report `Number of completed Surveys` as one of the metrics.  What is the `median` of the number of completed surveys per facility?

Note that some rows don't have a valid number in them for the number of completed surveys.


<!-- grade_id: q08-solution, points: 0, grade: false, solution: true -->
```python

x_variable = pd.DataFrame(hhcaps['Number of completed Surveys'])
y_variable = x_variable[x_variable['Number of completed Surveys'].str.isnumeric()]


answer = y_variable.agg('median')['Number of completed Surveys']
print(answer)
```

Output:
```
84.0
```

<!-- grade_id: q08-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == float or type(answer) == numpy.float64)
assert(answer == 84)

```

```python

```

## #09 - 

How many providers are there from St. Louis, Missouri?

*Note that the City columnd may have trailing spaces that need to be accomodated or cleaned up*

<!-- grade_id: q09-solution, points: 0, grade: false, solution: true -->
```python
def saint_louis(df, column):
    x_variable = (df[column].astype(str).strip() == "SAINT LOUIS")
    return pd.Series(x_variable)

hhcaps.apply(saint_louis,column = 'City')
```

Output:
```
pandas/_libs/index.pyx in pandas._libs.index.IndexEngine.get_loc()

pandas/_libs/index_class_helper.pxi in pandas._libs.index.Int64Engine._check_type()

KeyError: ('City', 'occurred at index State')
```

<!-- grade_id: q09-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 23)

```

```python

```

# #10 - 

Which state got the highest percentage of 'top box' scorings, using the `Percent of patients who gave their home health agency a rating of 9 or 10 on a scale from 0 (lowest) to 10 (highest)` field.

Ignore those rows where the data is **Not Available**

```python
field_one = 'Percent of patients who gave their home health agency a rating of 9 or 10 on a scale from 0 (lowest) to 10 (highest)'
field_two = 'Number of completed Surveys'

x_variable = hhcaps[hhcaps[field_one].str.isnumeric() & hhcaps[field_two].str.isnumeric()].astype({field_one: int, field_two: float})
x_variable['top_count'] = (x_variable[field_one]*x_variable[field_two]/100)
top_box = x_variable.groupby('State')['top_count',field_two].agg('sum') 
top_box['percent'] = top_box['top_count'] 
top_box2 = top_box['percent'] / top_box[field_two] * 100
top_box3 = top_box2.sort_values(ascending=False)
print(top_box3)                                                 
                                                  
                                                  
answer = top_box3.index[0]
```

Output:
```
State
MS    89.009587
WV    88.631481
LA    88.514367
KY    88.246596
AL    87.925021
SC    87.540108
ME    86.991524
AR    86.905304
NC    86.848742
PA    86.707218
TN    86.486643
VI    86.000000
MO    85.913440
DE    85.906279
TX    85.681190
SD    85.641256
MA    85.619339
OK    85.574008
IN    85.359125
IA    85.323734
RI    85.182714
ND    85.096143
NH    85.063360
KS    85.002900
GA    84.764847
IL    84.601804
WY    84.482639
VT    84.421548
OH    84.262428
WI    84.156338
ID    84.136851
PR    84.076573
NJ    83.947748
NE    83.922517
VA    83.887008
FL    83.852177
MI    83.773439
MD    83.669581
MT    83.488931
CT    83.348512
AK    83.332054
CO    83.253737
NM    83.166153
HI    82.966929
UT    81.896078
WA    81.683677
AZ    81.453258
MN    80.586314
CA    80.558193
OR    80.364508
NV    80.122497
NY    79.512880
MP    78.090909
DC    74.017812
GU    64.584615
dtype: float64
```

<!-- grade_id: q10-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'MS')
```

```python

```