# Week 2 Exercises

_McKinney 2.3_

**All of our homework will be structured a specific way to make it easier to verify for yourself that you're on the right track as well as easier for me to do a first pass on automated grading.  The structure of each question will require you to write a function using Python code.  Below each question are some tests (`assertions`) that verify your code is working correctly.  If any assertions fail, you know that something isnt' right with your code.**

**For now, please don't change any function names**

## Part 1: Multiple Choice

For the multiple choice questions below, type in the letter of the best answer where indicated.  Contrary to the note above, this section won't tell you if you have the correct answers or not immediately.

### 1.1 Code Blocks

Why is some code in Python indented from other lines of code?
> A) To specify that line of code should only be executed if the previous line is True
>
> B) The programmer can choose when to indent code to make it easier to read
>
> C) To block code apart to designate it should only be executed at certain times
>
> D) Indented lines are treated like comments and not executed

```python
answer11 = "X"
```

### 1.2 Python Versions

Which statements about Python 2 and Python 3 are true?
> A) Python version 2 and version 3 have completely compatible syntax
>
> B) Any code that runs in Python version 3 will run in version 2
> 
> C) Python version 3 and version 2 packages are compatible
> 
> D) Python version 2 and version 3 have some differences, but are mostly the same

```python
answer12 = "X"
```

## 1.3 Python Syntax

Which other common programming languages does Python look like?
> A) C++
>
> B) Java
> 
> C) JavaScript
>
> D) None of these

```python
answer13 = "X"
```

### 2.1 Variables
In the scenario below, which statments about the variables `a` and `b` are true.
```
a = [1, 2, 3]
b = a
```

> A) `a` is a list and `b` is a reference to another variables
> 
> B) `a` and `b` both refer to a list that has three elements
> 
> C) `b` is a copy of `a`
>
> D) `b` refers to `a` which refers to a list with three elements


```python
answer21 = "X"
```

### 2.2 Variable Names

Which of the following is **NOT** true about how you should name variables?

> A) Variable names should be as short as possible
> 
> B) Variable names can use abbreviations
>
> C) Variable names should be readable by humans
>
> D) Variable names represent a kind of documentation

```python
answer22 = "X"
```

### 2.3 Why Use Variables

Based on the purpose and best practices for using variables in Python, which of the following lines of code represents the best use of variables?  The code in the example below is used to convert degrees Farenheit to degress Celsius.

> A) `v2 = (v1 - 32) * 5/9`
>
> D) `temp_c = (temp_f - 32) * 5/9`
>
> B) `tempcelsius = (tempfarenheit - 32) * 5/9`
>
> C) `temp-celsius = (temp-farenheight - 32) * 5/9`

```python
answer23 = "X"
```

### 3.1 Built-in Types

Which of the following are NOT standard built-in data types in Python.  Put your answer inside brackets as shown here: `answer = ["X", "Y", "Z"]` 

> A) bool
>
> B) datetime
>
> C) decimal
>
> D) float
>
> E) int
>
> F) timestamp
>
> G) None


```python
answer31 = ["X","X"]
```

### 3.2 Strings

Which is the proper way to create a multi-line string in a Python program?

> A) Place `#` characters at the beginning of each line of the string
>
> B) Start the string with a `'` or `"` and then indent all of the subsquent lines until you place a `'` or `"` at the end of the multi-line string
>
> C) Place triple-quotes (`'''` or `"""`) on the first line and last line of the multi-line string
>
> D) You can't create multi-line strings in Python

```python
answer32 = "X"
```

### 3.3 Combining Data

Which of the following is not a valid operation?

> A) `'This is a ' + "complete sentence."`
>
> B) `'She is ' + 93 + ' years old.'`
>
> C) `"Repeat this! " * 4`
>
> D) `3.14 * 7`

```python
answer33 = "X"
```

### 4.1 Math

What will be the output of the following lines of code?  (Yes, I know that you can easily copy and paste this into another notebook and run it.  Please try to reason through it!)

```python
temp_c = 100
temp_f = temp_c * 9//5 + 32
temp_f
```

> A) 212
>
> B) 132
>
> C) 32
> 
> D) 98.6

```python
answer41 = "X"
```

Which code represents the proper way to represent the mathematical expression $pi^2$ ?

> A) `pi^2`
>
> B) `pi x pi`
>
> C) `pi ** 2`
>
> D| `pi | 2`

```python
answer42 = "X"
```

### 5.1 Averaging three numbers

Image that you've got a series of body temperature measurements and we want to find the average of three of them.  Write code below that computes the average of three numbers represented by the variables `temp1`, `temp2`, `temp3` and assigns the result to the variables `avg_temp`.

<!-- grade_id: average-answer, points: 0, grade: false, solution: true -->
```python
def average_temp(temp1, temp2, temp3):
    
    ### BEGIN SOLUTION

    ### END SOLUTION
    
    return avg_temp
```

Your function should return 98 if you give it inputs of 96, 99, 99.  Check to see that it does by typeing `average_temp(96, 99, 101)` in the cell below and running it.

```python

```

<!-- grade_id: average-tests, points: 5, grade: true, solution: false -->
```python
assert average_temp(96, 99, 99) == 98
assert average_temp(0, 0, 0) == 0
assert average_temp(100, -100, 0) == 0
assert average_temp(96, 98, 100) == 98
assert average_temp(96.5, 98.5, 100.5) == 98.5
assert average_temp(96, 98, 99) == (293 / 3)
```

### 5.2 Calculation

Create new names for the variables that you need in the code framework below and use the variables to compute the corresponding body mass index (BMI). You can use any meaningful variable names you like. When you compute the BMI, assign a variable to it, and then print the result to the notebook's output.

* Weght in kilograms is 90.7
* Height in centimeters in 182

Note that a simple BMI calculation is weight in kilograms divided by the square of height in meters.

For this problem, you'll be graded both on proper calculation as well as on programming style and appropriate variables names.

_Also note that BMI should not be considered a universal measure of health or fitness. A health BMI can vary greatly depending on many other fitness conditions and body types.

<!-- grade_id: bmi-answer, points: 0, grade: false, solution: true -->
```python
def bmi(parameter1, parameter2):
    
    # 1. Change the names of parameter1 and parameter2 to be something more meaningful
    # 2. Change the name of the return variable below
    
    ### BEGIN SOLUTION

    ### END SOLUTION
    
    return return_value
```

Test your code below.  The BMI for a weight of 90.7 kg and height of 182 cm should be approximately 27.38

```python

```

<!-- grade_id: bmi-test, points: 0, grade: true, solution: false -->
```python
assert round(bmi(90.7, 182), 2) == 27.38
assert round(bmi(91, 180), 2) == 28.09
assert round(bmi(70, 170), 2) == 24.22
```

<!-- grade_id: bmi-style, points: 5, grade: true, solution: true -->
```python

```

### 5.3 BMI
Assuming that someone doesn't change height over time, what is the approximate number of days it will take for someone to get from the BMI you found in #3 above to a BMI of 24.9 (typically regarded as being in a healthy range) while losing weight at a rate of 0.75 kg per week.

As needed, use variables and multiple calculation steps to make your thought process visible in the code. For instance, rather than using the notebook as a calculator, use variables for intermediate answers.

Again, you will receive points both for correct answers as well as for style.


<!-- grade_id: weight-answer, points: 0, grade: false, solution: true -->
```python
def lose_weight(weight, height, kg_per_week, target_bmi):
    
    # Note that the variables you need area already named for you:
    # height is the height in cm
    # weight is the weight in kg
    # kg_per_week is how much the person wants to lose per week
    # target_bmi is the intended target bmi
    
    ### BEGIN SOLUTION

    ### END SOLUTION
    
    return days
```

Test your code below.  It should return approximately 76.73 days for the given parameters above.

```python
lose_weight(70, 170, 2, 22)
```

<!-- grade_id: weight-tests, points: 5, grade: true, solution: false -->
```python
assert round(lose_weight(90.7, 182, 0.75, 24.9),2) == 76.73
assert round(lose_weight(91, 180, 0.75, 24.9),2) == 96.36
assert round(lose_weight(70, 170, 0.75, 22),2) == 59.92
assert round(lose_weight(70, 170, 2, 22),2) == 22.47

```

<!-- grade_id: weight-style, points: 5, grade: true, solution: true -->
```python

```